# Unit tests



```python
from nose.tools import assert_equal
from nose.tools import assert_almost_equal
```

## Problem 1

The main purpose of this problem is to learn how to assign values to a variable.

In the following code cell, assign a value of 141 to an integer variable named `the_integer`.

<!-- grade_id: int, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
the_integer = 141
```

The two assert statements below are called unit tests. They are important to check the values of the variables and their data types. 
Writing the tests first, and then writing the actual code is actually a perfectly valid approach (preferred by many) to programming. It's called test-driven development.



<!-- grade_id: int-test, points: 5, grade: true, solution: false -->
```python
assert_equal(the_integer,141)
assert_equal(type(the_integer), int)
```

In the following code cell, assign the value 1.6180339 to a floating point variable named `floating_bubbles`.

<!-- grade_id: float, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
floating_bubbles = 3.145679
```

<!-- grade_id: float-test, points: 5, grade: true, solution: false -->
```python
assert_almost_equal(floating_bubbles, 3.145679)
assert_equal(type(floating_bubbles), float)
```

Finally, you will need to assign the string 'YoImAString' to a string variable named `the_string`.

<!-- grade_id: string, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
the_string = 'YoImAString'
```

<!-- grade_id: string-test, points: 10, grade: true, solution: false -->
```python
assert_equal(the_string, 'YoImAString')
assert_equal(type(the_string), str)
```

```python

```